---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: cbbdeb8cf4
xp: 50
```

`@projector_key`
75b2e12632fcf88c300c505d0b1f6a5c

---

## Lesson 2.1 - Relational to Flat structure

```yaml
type: VideoExercise
key: 0a0d710add
xp: 50
```

`@projector_key`
b7100554f028184c77b6c33abd78e2d5
